Overflow metabolism originates from growth optimization and cell heterogeneity

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public Review):

Summary:

Cell metabolism exhibits a well-known behavior in fast-growing cells, which employ seemingly wasteful fermentation to generate energy even in the presence of sufficient environmental oxygen. This phenomenon is known as Overflow Metabolism or the Warburg effect in cancer. It is present in a wide range of organisms, from bacteria and fungi to mammalian cells.

In this work, starting with a metabolic network for Escherichia coli based on sets of carbon sources, and using a corresponding coarse-grained model, the author applies some well-based approximations from the literature and algebraic manipulations. These are used to successfully explain the origins of Overflow Metabolism, both qualitatively and quantitatively, by comparing the results with E. coli experimental data.

By modeling the proteome energy efficiencies for respiration and fermentation, the study shows that these parameters are dependent on the carbon source quality constants K_i (p.115 and 116). It is demonstrated that as the environment becomes richer, the optimal solution for proteome energy efficiency shifts from respiration to fermentation. This shift occurs at a critical parameter value K_A(C).

This counter intuitive results qualitatively explains Overflow Metabolism.

Quantitative agreement is achieved through the analysis of the heterogeneity of the metabolic status within a cell population. By introducing heterogeneity, the critical growth rate is assumed to follow a Gaussian distribution over the cell population, resulting in accordance with experimental data for E. coli. Overflow metabolism is explained by considering optimal protein allocation and cell heterogeneity.

The obtained model is extensively tested through perturbations: 1) Introduction of overexpression of useless proteins; 2) Studying energy dissipation; 3) Analysis of the impact of translation inhibition with different sub-lethal doses of chloramphenicol on Escherichia coli; 4) Alteration of nutrient categories of carbon sources using pyruvate. All model perturbations results are corroborated by E. coli experimental results.

Strengths:

In this work, the author effectively uses modeling techniques typical of Physics to address complex problems in Biology, demonstrating the potential of interdisciplinary approaches to yield novel insights. The use of Escherichia coli as a model organism ensures that the assumptions and approximations are well-supported in existing literature. The model is convincingly constructed and aligns well with experimental data, lending credibility to the findings. In this version, the extension of results from bacteria to yeast and cancer is substantiated by a literature base, suggesting that these findings may have broad implications for understanding diverse biological systems.

We appreciate the reviewer’s exceptionally positive comments. The manuscript has been significantly improved thanks to the reviewer’s insightful suggestions.

Weaknesses:

The author explores the generalization of their results from bacteria to cancer cells and yeast, adapting the metabolic network and coarse-grained model accordingly. In previous version this generalization was not completely supported by references and data from the literature. This drawback, however, has been treated in this current version, where the authors discuss in much more detail and give references supporting this generalization.

We appreciate the reviewer’s recognition of our revisions and the insightful suggestions provided in the previous round, which have greatly strengthened our manuscript.

Reviewer #2 (Public Review):

In this version of manuscript, the author clarified many details and rewrote some sections. This substantially improved the readability of the paper. I also recognized that the author spent substantial efforts in the Appendix to answer the potential questions.

We thank the reviewer for the positive comments and the suggestions to improve our manuscript.

Unfortunately, I am not currently convinced by the theory proposed in this paper. In the next section, I will first recap the logic of the author and explain why I am not convinced. Although the theory fits many experimental results, other theories on overflow metabolism are also supported by experiments. Hence, I do not think based on experimental data we could rule in or rule out different theories.

We thank the reviewer for both the critical and constructive comments.

Regarding the comments on the comparison between theoretical and experimental results, we would like to first emphasize that no prior theory has resolved the conflict arising from the proteome efficiencies measured in E. coli and eukaryotic cells. Specifically, prevalent explanations (Basan et al., Nature 528, 99–104 (2015); Chen and Nielsen, PNAS 116, 17592–17597 (2019)) hold that overflow metabolism results from proteome efficiency in fermentation consistently being higher than that in respiration. While it was observed in E. coli that proteome efficiency in fermentation exceeds that in respiration when cells were cultured in lactose at saturated concentrations (Basan et al., Nature 528, 99-104 (2015)), more recent findings (Shen et al., Nature Chemical Biology 20, 1123–1132 (2024)) show that the measured proteome efficiency in respiration is actually higher than in fermentation for many yeast and cancer cells, despite the presence of aerobic glycolytic fermentation flux. To the best of our knowledge, no prior theory has explained these contradictory experimental results. Notably, our theory resolves this conflict and quantitatively explains both sets of experimental observations (Basan et al., Nature 528, 99-104 (2015); Shen et al., Nature Chemical Biology 20, 1123–1132 (2024)) by incorporating cell heterogeneity and optimizing cell growth rate through protein allocation.

Furthermore, rather than merely fitting the experimental results, as explained in Appendices 6.2, 8.1-8.2 and summarized in Appendix-tables 1-3, nearly all model parameters important for our theoretical predictions for E. coli were derived from in vivo and in vitro biochemical data reported in the experimental literature. For comparisons between model predictions and experimental results for yeast and cancer cells (Shen et al., Nature Chemical Biology 20, 1123–1132 (2024)), we intentionally derived Eq. 6 to ensure an unbiased comparison.

Finally, in response to the reviewer’s suggestion, we have revised the expressions in our manuscript to present the differences between our theory and previous theories in a more modest style.

Recap: To explain the origin of overflow metabolism, the author uses the following logic:

(1) There is a substantial variability of single-cell growth rate

(2) The flux (J_r^E) and (J_f^E) are coupled with growth rate by Eq. 3

(3) Since growth rate varies from cells to cells, flux (J_r^E) and (J_f^E) also varies (4) The variabilities of above fluxes in above create threshold-analog relation, and hence overflow metabolism.

We thank the reviewer for the clear summary. We apologize for not explaining some points clearly enough in the previous version of our manuscript, which may have led to misunderstandings. We have now revised the relevant content in the manuscript to clarify our reasoning. Specifically, we have applied the following logic in our explanation:

(a) The solution for the optimal growth strategy of a cell under a given nutrient condition is a binary choice between respiration and fermentation, driven by comparing their proteome efficiencies (εr and εf ).

(b) Under nutrient-poor conditions, the nutrient quality (κA) is low, resulting in the proteome efficiency of respiration being higher than that of fermentation (i.e., εr > εf), so the cell exclusively uses respiration.

(c) In rich media (with high κA), the proteome efficiency of fermentation increases more rapidly and surpasses that of respiration (i.e., εf > εr ), hence the cell switches to fermentation.

(d) Heterogeneity is introduced: variability in the κcat of catalytic enzymes from cell to cell. This leads to heterogeneity (variability) in εr and εf within a population of cells under the same nutrient condition.

(e) The critical value of nutrient quality for the switching point (, where εr= εf ) changes from a single point to a distribution due to cell heterogeneity. This results in a distribution of the critical growth rate λC (defined as ) within the cell population.

(f) The change in culturing conditions (with a highly diverse range of κA) and heterogeneity in the critical growth rate λC (a distribution of values) result in the threshold-analog relation of overflow metabolism at the cell population level.

Steps (a)-(c) were applied to qualitatively explain the origin of overflow metabolism, while steps (d)-(f) were further used to quantitatively explain the threshold-analog relation observed in the data on overflow metabolism.

Regarding the reviewer’s recap, which seems to have involved some misunderstandings, we first emphasize that the major change in cell growth rate for the threshold-analog relation of overflow metabolism—particularly as it pertains to logic steps (1), (3) and (4)—is driven by the highly varied range of nutrient quality (κA) in the culturing conditions, rather than by heterogeneity between cells. For the batch culture data, the nutrient type of the carbon source differs significantly (e.g., Fig.1 in Basan et al., Nature 528, 99-104 (2015), wild-type strains). In contrast, for the chemostat data, the concentration of the carbon source varies greatly due to the highly varied dilution rate (e.g., Table 7 in Holms, FEMS Microbiology Reviews 19, 85-116 (1996)). Both of these factors related to nutrient conditions are the major causes of the changes in cell growth rate in the threshold-analog relation.

Second, Eq. 3, as mentioned in logic step (2), represents a constraint between the fluxes ( and ) and the growth rate (λ) for a single nutrient condition (with a given value of κA ideally) rather than for varied nutrient conditions. For a single cell in each nutrient condition, the optimal growth strategy is binary, between respiration and fermentation.

Finally, for the threshold-analog relation of overflow metabolism, the switch from respiration to fermentation is caused by the increased nutrient quality in the culturing conditions, rather than by cell heterogeneity as indicated in logic step (4). Upon nutrient upshifts, the proteome efficiency of fermentation surpasses that of respiration, causing the optimal growth strategy for the cell to switch from respiration to fermentation. The role of cell heterogeneity is to transform the growth rate-dependent fermentation flux in overflow metabolism from a digital response to a threshold-analog relation under varying nutrient conditions.

My opinion:

The logic step (2) and (3) have caveats. The variability of growth rate has large components of cellular noise and external noise. Therefore, variability of growth rate is far from 100% correlated with variability of flux (J_r^E) and (J_f^E) at the single-cell level. Single-cell growth rate is a complex, multivariate functional, including (Jr^E) and (J_f^E) but also many other variables. My feeling is the correlation could be too low to support the logic here.

One example: ribosomal concentration is known to be an important factor of growth rate in bulk culture. However, the "growth law" from bulk culture cannot directly translate into the growth law at single-cell level [Ref1,2]. This is likely due to other factors (such as cell aging, other muti-stability of cellular states) are involved.

Therefore, I think using Eq.3 to invert the distribution of growth rate into the distribution of (Jr^E) and (J_f^E) is inapplicable, due to the potentially low correlation at single-cell level. It may show partial correlations, but may not be strong enough to support the claim and create fermentation at macroscopic scale.

Overall, if we track the logic flow, this theory implies overflow metabolism is originated from variability of k_cat of catalytic enzymes from cells to cells. That is, the author proposed that overflow metabolism happens macroscopically as if it is some "aberrant activation of fermentation pathway" at the single-cell level, due to some unknown partially correlation from growth rate variability.

We thank the reviewer for raising these questions and for the insights. We apologize for any lack of clarity in the previous version of our manuscript that may have caused misunderstandings. We have revised the manuscript to address all points, and below are our responses to the questions, some of which seem to involve misunderstandings.

First, in our theory, the qualitative behavior of overflow metabolism—where cells use respiration under nutrient-poor conditions (low growth rate) and fermentation in rich media (high growth rate)—does not arise from variability between cells, as the reviewer seems to have interpreted. Instead, it originates from growth optimization through optimal protein allocation under significantly different nutrient conditions. Specifically, the proteome efficiency of fermentation is lower than that of respiration (i.e. εf < εr) under nutrient-poor conditions, making respiration the optimal strategy in this case. However, in rich media, the proteome efficiency of fermentation surpasses that of respiration (i.e. εf < εr), leading the cell to switch to fermentation for growth optimization. To implement the optimal strategy, as clarified in the revised manuscript and discussed in Appendix 2.4, a cell should sense and compare the proteome efficiencies between respiration and fermentation, choosing the pathway with the higher efficiency, rather than sensing the growth rate, which can fluctuate due to stochasticity. Regarding the role of cell heterogeneity in overflow metabolism, as discussed in our previous response, it is twofold: first, it quantitatively illustrates the threshold-analog response of growth rate-dependent fermentation flux, which would otherwise be a digital response without heterogeneity during growth optimization; second, it enables us to resolve the paradox in proteome efficiencies observed in E. coli and eukaryotic cells, as raised by Shen et al. (Shen et al., Nature Chemical Biology 20, 1123–1132 (2024)).

Second, regarding logic step (2) in the recap, the reviewer thought we had coupled the growth rate (λ) with the respiration and fermentation fluxes ( and ) through Eq. 3, and used Eq. 3 to invert the distribution of growth rate into the distribution of respiration and fermentation fluxes. We need to clarify that Eq. 3 represents the constraint between the fluxes and the growth rate under a single nutrient condition, rather than describing the relation between growth rate and the fluxes ( and κA), without considering optimal protein allocation, the cell growth rate varies with the fluxes according to Eq.3 by adjusting the proteome allocation between respiration and fermentation (ϕr and ) under varied nutrient conditions. In a given nutrient condition (with a fixed value of ϕf). However, once growth optimization is applied, the optimal protein allocation strategy for a cell is limited to either pure respiration (with ϕf =0 and ϕr =0 and κA). Furthermore, under varying nutrient conditions (with different values of ) or pure fermentation (with κA), both proteome efficiencies of respiration and fermentation (εr and (εf) change with nutrient quality ), depending on the nutrient condition (or the value of κA (see Eq. 4). Thus, Eq. 3 does not describe the relation between growth rate (λ) and the fluxes ( and ) under nutrient variations.

Thirdly, regarding reviewer’s concerns on logic step (3) in the recap, as well as the example where ribosome concentration does not correlate well with cell growth rate at the single-cell level, we fully agree with reviewer that, due to factors such as stochasticity and cell cycle status, the growth rate fluctuates constantly for each cell. Consequently, it would not be fully correlated with cell parameters such as ribosome concentration or respiration/fermentation flux. We apologize for our oversight in not discussing suboptimal growth conditions in the previous version of the manuscript. In response, we have added a paragraph to the discussion section and a new Appendix 2.4, titled “Dependence of the model on optimization principles,” to address these issues in detail. Specifically, recent experimental studies (Dai et al., Nature microbiology 2, 16231 (2017); Li et al., Nature microbiology 3, 939–947 (2018)) show that the inactive portion of ribosomes (i.e., ribosomes not bound to mRNAs) can vary under different culturing conditions. The reviewer also pointed out that ribosome concentration does not correlate well with cell growth rate at single-cell level. In this regard, we have cited Pavlou et al. (Pavlou et al., Nature Communications 16, 285 (2025)) instead of the references provided by the reviewer (Ref1 and Ref2), with our rationale outlined in the final section of the author response. These findings (Dai et al, (2017); Li et al., (2018); Pavlou et al., (2025)) suggest that ribosome allocation may be suboptimal under many culturing conditions, likely as cells prepare for potential environmental changes (Li et al., Nature microbiology 3, 939–947 (2018)). However, since our model's predictions regarding the binary choice between respiration and fermentation are based solely on comparing proteome efficiency between these two pathways, the optimal growth principle in our model can be relaxed. Specifically, efficient protein allocation is required only for enzymes rather than ribosomes, allowing our model to remain applicable under suboptimal growth conditions. Furthermore, protein allocation via the ribosome occurs at the single-cell level rather than at the population level. The strong linear correlation between ribosomal concentration and growth rate at the population level under nutrient variations suggests that each cell optimizes its protein allocation individually. Therefore, the principle of growth optimization still applies to individual cells, although factors like stochasticity, nutrient variation preparations, and differences in cell cycle stages may complicate this relationship, resulting in only a rough linear correlation between ribosome concentration and growth rate at the single-cell level (with with R2 = 0.64 reported in Pavlou et al., (2025)).

Lastly, regarding the reviewer concerns about the heterogeneity of fermentation and respiration at macroscopic scale, we first clarify in the second paragraph of this response that the primary driving force for cells to switch from respiration to fermentation in the context of overflow metabolism is the increased nutrient quality under varying culturing conditions, which causes the proteome efficiency of fermentation to surpass that of respiration. Under nutrient-poor conditions, our model predicts that all cells use respiration, and therefore no heterogeneity for the phenotype of respiration and fermentation arises in these conditions. However, in a richer medium, particularly one that does not provide optimal conditions but allows for an intermediate growth rate, our model predicts that some cells opt for fermentation while others continue with respiration due to cell heterogeneity (with εf > εr for some cells engaging in fermentation and εr > εf for the other cells engaging in respiration within the same medium). Both of these predictions have been validated in isogenic singlecell experiments with E. coli (Nikolic et al., BMC Microbiology 13, 258 (2013)) and S. cerevisiae (Bagamery et al., Current Biology 30, 4563–4578 (2020)). The single-cell experiments by Nikolic et al. with E. coli in a rich medium of intermediate growth rate clearly show a bimodal distribution in the expression of genes related to overflow metabolism (see Fig. 5 in Nikolic et al., BMC Microbiology 13, 258 (2013)), where one subpopulation suggests purely fermentation, while the other suggests purely respiration. In contrast, in a medium with lower nutrient concentration (and consequently lower nutrient quality), only the respirative population exists (see Fig. 5 in Nikolic et al., BMC Microbiology 13, 258 (2013)). These experimental results from E. coli (Nikolic et al., BMC Microbiology 13, 258 (2013)) are fully consistent with our model predictions. Similarly, the single-cell experiments with S. cerevisiae by Bagamery et al. clearly identified two subpopulations of cells with respect to fermentation and respiration in a rich medium, which also align well with our model predictions regarding heterogeneity in fermentation and respiration within a cell population in the same medium.

Compared with other theories, this theory does not involve any regulatory mechanism and can be regarded as a "neutral theory". I am looking forward to seeing single cell experiments in the future to provide evidences about this theory.

We thank the reviewer for raising these questions and for the valuable insights. Regarding the regulatory mechanism, we have now added a paragraph in the discussion section of our manuscript and Appendix 2.4 to address this point. Specifically, our model predicts that a cell can implement the optimal strategy by directly sensing and comparing the proteome efficiencies of respiration and fermentation, choosing the pathway with the higher efficiency. At the gene regulatory level, a growing body of evidence suggests that the cAMP-CRP system plays an important role in sensing and executing the optimal strategy between respiration and fermentation (Basan et al., Nature 528, 99-104 (2015); Towbin et al., Nature Communications 8, 14123 (2017); Valgepea et al., BMC Systems Biology 4, 166 (2010); Wehrens et al., Cell Reports 42, 113284 (2023)). However, it has also been suggested that the cAMP-CRP system alone is insufficient, and additional regulators may need to be identified to fully elucidate this mechanism (Basan et al., Nature 528, 99-104 (2015); Valgepea et al., BMC Systems Biology 4, 166 (2010)).

Regarding the single-cell experiments that provide evidence for this theory, we have shown in the previous paragraphs of this response that the heterogeneity between respiration and fermentation, as predicted by our model for isogenic cells within the same culturing condition, has been fully validated by single-cell experiments with E. coli (Fig. 5 from Nikolic et al., BMC Microbiology 13, 258 (2013)) and S. cerevisiae (Fig. 1 and the graphical abstract from Bagamery et al., Current Biology 30, 4563–4578 (2020)). We have now revised the discussion section of our manuscript to make this point clearer.

[Ref1] https://www.biorxiv.org/content/10.1101/2024.04.19.590370v2

[Ref2] https://www.biorxiv.org/content/10.1101/2024.10.08.617237v2

We thank the reviewer for providing insightful references. Regarding the two specific references, Ref1 directly addresses the deviation in the linear relationship between growth rate and ribosome concentration (“growth law”) at the single-cell level. However, since the authors of Ref1 determined the rRNA abundance in each cell by aligning sequencing reads to the genome, this method inevitably introduces a substantial amount of measurement noise. As a result, we chose not to cite or discuss this preprint in our manuscript. Ref2 appears to pertain to a different topic, which we suspect may be a copy/paste error. Based on the reviewer’s description and the references in Ref1, we believe the correct Ref2 should be Pavlou et al., Nature Communications 16, 285 (2025) (with the biorxiv preprint link: https://www.biorxiv.org/content/10.1101/2024.04.26.591328v1). In this reference, it is stated that the relationship between ribosome concentration and growth rate only roughly aligns with the “growth law” at the single-cell level (with R2 = 0.64), exhibiting a certain degree of deviation. We have now cited and incorporated the findings of Pavlou et al. (Pavlou et al., Nature Communications 16, 285 (2025)) in both the discussion section of our manuscript and Appendix 2.4. Overall, we agree with Pavlou et al.’s experimental results, which suggest that ribosome concentration does not exhibit a strong linear correlation with cell growth rate at the single-cell level. However, we remain somewhat uncertain about the extent of this deviation, as Pavlou et al.’s experimental setup involved alternating nutrients between acetate and glucose, and the lapse of five generations may not have been long enough for the growth to be considered balanced. Furthermore, as observed in Supplementary Movie 1 of Pavlou et al., some of the experimental cells appeared to experience growth limitations due to squeezing pressure from the pipe wall of the mother machine, which could further increase the deviation from the “growth law” at the single-cell level.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

I have no specific comments for the authors related to this last version of the paper. I believe the authors have properly improved the previous version of the manuscript.

Response: We thank the reviewer for the highly positive comments and for recognizing the improvements made in the revised version of our manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

A classic problem in metabolism is that fast-proliferating cells use seemingly wasteful fermentation for energy biogenesis in the presence of sufficient oxygen. This counterintuitive phenomenon, known as overflow metabolism or the Warburg effect, is universal across various organisms. Despite extensive research, its origin and function remain unclear. Here, we show that overflow metabolism can be understood through growth optimization combined with cell heterogeneity. A model of optimal protein allocation, coupled with heterogeneity in enzyme catalytic rates among cells, quantitatively explains why and how cells choose between respiration and fermentation under different nutrient conditions. Our model quantitatively illustrates the growth rate dependence of fermentation flux and enzyme allocation under various perturbations and is fully validated by experimental results in Escherichia coli. Our work provides a quantitative explanation for the Crabtree effect in yeast and the Warburg effect in cancer cells and can be broadly used to address heterogeneity-related challenges in metabolism.

Introduction

A prominent feature of cancer metabolism is that tumor cells excrete large quantities of fermentation products in the presence of sufficient oxygen (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). This process, discovered by Otto Warburg in the 1920s (Warburg et al., 1924) and known as the Warburg effect, aerobic glycolysis, or overflow metabolism (Basan et al., 2015; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009), is ubiquitous among fast-proliferating cells across a broad spectrum of organisms (Vander Heiden et al., 2009), ranging from bacteria (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; Neidhardt et al., 1990) and fungi (De Deken, 1966) to mammalian cells (Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). For microbes, cells use standard respiration when nutrients are scarce, while they use the counterintuitive aerobic glycolysis when nutrients are adequate, just analogous to normal tissues and cancer cells, respectively (Vander Heiden et al., 2009).

Earlier version: Over the past century, and particularly through extensive studies in the last two decades (Liberti and Locasale, 2016), various rationales for overflow metabolism have been proposed (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011). Notably, Basan et al. (Basan et al., 2015) provided a systematic characterization of this process, including various types of experimental perturbations. Currently, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) hold that overflow metabolism arises from the proteome efficiency in fermentation being consistently higher than that in respiration. However, recent studies have shown that the measured proteome efficiency in respiration is actually higher than in fermentation for many yeast and cancer cells (Shen et al., 2024), even though these cells generate fermentation products through aerobic glycolysis. This finding (Shen et al., 2024) apparently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Furthermore, most explanations (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Niebel et al., 2019; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011) heavily rely on the assumption that cells optimize their growth rate for a given rate of carbon influx (i.e., nutrient uptake rate) under each nutrient condition (or its equivalents). However, this assumption is questionable, as the given factors in a nutrient condition are the identities and concentrations of the carbon sources (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the carbon influx. Therefore, the origin and function of overflow metabolism still remain unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009).

Revised version: Over the past century, and particularly through extensive studies in the last two decades (Liberti and Locasale, 2016), various rationales for overflow metabolism have been proposed (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011). Notably, Basan et al. (Basan et al., 2015) provided a systematic characterization of this process, including various types of experimental perturbations. Currently, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) hold that overflow metabolism arises from the proteome efficiency in fermentation being consistently higher than that in respiration. However, recent studies have shown that the measured proteome efficiency in respiration is actually higher than in fermentation for many yeast and cancer cells (Shen et al., 2024), even though these cells generate fermentation products through aerobic glycolysis. This finding (Shen et al., 2024) apparently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Furthermore, most explanations (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011) rely on the assumption that cells optimize their growth rate for a given rate of carbon influx (i.e., nutrient uptake rate) under each nutrient condition (or its equivalents). However, this assumption remains open to further scrutiny, as the given factors in a nutrient condition are the identities and concentrations of the carbon sources (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the carbon influx. Therefore, the origin and function of overflow metabolism still remain unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009).

Why have microbes and cancer cells evolved to possess the seemingly wasteful strategy of aerobic glycolysis? For unicellular organisms, there is evolutionary pressure (Vander Heiden et al., 2009) to optimize cellular resources for rapid growth (Dekel and Alon, 2005; Edwards et al., 2001; Hui et al., 2015; Li et al., 2018; Scott et al., 2010; Towbin et al., 2017; Wang et al., 2019; You et al., 2013). In particular, it has been shown that cells allocate protein resources for optimal growth (Hui et al., 2015; Scott et al., 2010; Wang et al., 2019; You et al., 2013), and the most efficient protein allocation corresponds to elementary flux mode (Müller et al., 2014; Wortel et al., 2014). For cancer cells, disrupting the growth control system and evading immune destruction from the host are prominent hallmarks of their survival (Hanahan and Weinberg, 2011), which in certain ways mimic the evolutionary pressure on microbes to optimize cell growth rate. In this study, we apply the optimal growth principle of microbes, which also roughly holds for cancer cells, to a heterogeneous framework to address the puzzle of aerobic glycolysis. We use Escherichia coli as a typical example to show that overflow metabolism can be understood from optimal protein allocation combined with heterogeneity in enzyme catalytic rates. The optimal growth strategy varies between respiration and fermentation depending on the concentration and type of the nutrient, and the combination with cell heterogeneity results in the standard picture (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) of overflow metabolism. Our model quantitatively illustrates the growth rate dependence of fermentation/respiration flux and enzyme allocation under various types of perturbations in E. coli. Furthermore, it provides a quantitative explanation for the data on the Crabtree effect in yeast and the Warburg effect in cancer cells (Bartman et al., 2023; Shen et al., 2024).

Results

Coarse-grained model

Based on the topology of the metabolic network (Neidhardt et al., 1990; Nelson et al., 2008) (see Fig. 1A), we classify the carbon sources that enter from the upper part of glycolysis into Group A (Wang et al., 2019) and the precursors of biomass components (such as amino acids) into five pools. Specifically, each pool is designated according to its entry point (see Fig. 1A and Appendix 1.2 for details): a1 (entry point: G6P/F6P), a2 (entry point: GA3P/3PG/PEP), b (entry point: pyruvate/acetyl-CoA), c (entry point: a-ketoglutarate), and d (entry point: oxaloacetate). Pools al and a2 are also combined as Pool a due to the joint synthesis of precursors. Then, the metabolic network for Group A carbon source utilization (see Fig. 1A) can be coarse-grained into a model shown in Fig. 1B (see Appendix 2.1 for details), where node A represents an arbitrary carbon source of Group A. Evidently, Fig. 1B is topologically identical to Fig. 1A. Each coarse-grained arrow in Fig. 1B represents a stoichiometric flux Ji, which delivers carbon flux and may be accompanied by energy consumption or biogenesis (e.g., J1, Ja1, see Figs. 1A-B and Appendix-fig. 1A).

Model and results of overflow metabolism in E. coli.

(A) The central metabolic network of carbon source utilization. The Group A carbon sources (Wang et al., 2019) are labeled with green squares. (B) Coarse-grained model for Group A carbon source utilization. (C) Model predictions (see Eqs. S47 and S160) and experimental results (Basan et al., 2015; Holms, 1996) of overflow metabolism, covering the data for all the Group A carbon sources shown in (A). (D) Growth rate dependence of respiration and fermentation fluxes (see Eqs. S47 and S160). (E) The proteome efficiencies for energy biogenesis in the respiration and fermentation pathways vary with growth rate as functions of the substrate quality of a Group A carbon source (see Eqs. S31 and S36). See Appendices 8 and 10 for model parameter settings and experimental data sources (Basan et al., 2015; Holms, 1996; Hui et al., 2015) for Figs. 1–4 of E. coli.

In fact, the stoichiometric flux Ji scales with the cell population. For comparison with experiments, we define the normalized flux Ji(N)≡Ji⋅m0/Mcarbon , which can be regarded as the flux per unit of biomass (the superscript “(N)” stands for normalized; see Appendix 1.3–1.4 for details). Here, Mcarbon represents the carbon mass of the cell population, and m0 is the weighted average carbon mass of metabolite molecules at the entry of precursor pools (see Eq. S17). Then, the cell growth rate λ can be represented by the total outflow of the normalized fluxes: λ=∑ia1,a2,b,c,dJi(N) (see Appendix 1.4). The normalized fluxes of respiration and fermentation are Jr(N)≡J4(N) and Jf(N)≡J6(N) , respectively (see Fig. 1A-B). In practice, each Ji(N) is characterized by two quantities: the proteomic mass fraction ϕi of the enzyme dedicated to carrying the flux and the substrate quality κi, such that Ji(N)=ϕi⋅κi . We take the Michaelis-Menten form for the enzyme kinetics (Nelson et al., 2008), and then κi≡κi⋅[Si][Si]+Ki (see Eq. S12 and Appendix 1.4 for details), where [Si] is the concentration of substrate Si, and Ki is the Michaelis constant. For each intermediate node and reaction along the pathway (e.g., node M1 in Ja1), the substrate quality κi can be approximated as a constant (see Appendix 1.5): κi≡κi⋅[Si][Si]+Ki≈ki , where [Si] ≥ Ki generally holds true in bacteria (Bennett et al., 2009; Park et al., 2016). However, the nutrient quality κA is a variable that depends on the nutrient type and concentration of a Group A carbon source (see Eq. S27).

Generally, there are three independent fates for a Group A carbon source in the metabolic network (Chen and Nielsen, 2019): fermentation, respiration, and biomass generation (see Appendix-fig. 1C-E). Each draws a distinct proteome fraction of ϕf, ϕr and ϕBM, with no overlap between them (see Appendix 2.1). The net effect of the first two fates is energy biogenesis, while the last one generates precursors for biomass, accompanied by energy biogenesis. By applying the proteomic constraint (Scott et al., 2010) that there is a maximum fraction ϕmax for proteome allocation (ϕmax ≈ 0.48 (Scott et al., 2010)), we have: ϕf+ϕr+ϕBM=ϕmax. (1)

Earlier version: In fact, Eq. 1 is equivalent to ϕR+ϕA+∑j=16ϕj+∑ia1,a2,b,c,dϕi=ϕmax (see Appendix 2.1 for derivation details), where ϕR and ϕA represent the proteomic mass fractions of the active ribosome-affiliated proteins and the cargo proteins responsible for the uptake of the Group A carbon source, respectively. During cell proliferation, ribosomes serve as the factories for protein synthesis and are primarily composed of proteins (Neidhardt et al., 1990; Nelson et al., 2008), while other biomass components, such as RNA, are optimally produced (Kostinski and Reuveni, 2020) in accordance with the growth rate determined by protein synthesis. Thus, the cell growth rate is proportional to ϕR: λ = ϕR ° κt, where κt is a parameter set by the translation rate (Scott et al., 2010) (see Appendix 1.1 for details), which can be approximated as a constant within the growth rate range of interest (Dai et al., 2016).

Revised version: In fact, Eq. 1 is equivalent to ϕR+ϕA+∑j=16ϕj+∑ia1,a2,b,c,dϕi=ϕmax (see Appendix 2.1 for derivation details), where ϕR and ϕA represent the proteomic mass fractions of the active ribosome-affiliated proteins and the cargo proteins responsible for the uptake of the Group A carbon source, respectively. During cell proliferation, ribosomes serve as the factories for protein synthesis and are primarily composed of proteins (Neidhardt et al., 1990; Nelson et al., 2008), while other biomass components, such as RNA, are optimally produced (Kostinski and Reuveni, 2020) in accordance with the growth rate determined by protein synthesis. Thus, the cell growth rate is proportional to ϕR: λ = ϕR · κt, where κt is a parameter set by the translation rate (Scott et al., 2010) (see Appendix 1.1 for details), which can be approximated as a constant within the growth rate range of interest (Dai et al., 2017).

For balanced cell growth in bacteria, the energy demand JE, expressed as the stoichiometric energy flux in ATP, is generally proportional to the biomass production rate (Ebenhoh et al., 2024), since the proportion of maintenance energy is roughly negligible (Locasale and Cantley, 2010) (see Appendix 9 for the cases of yeast and tumor cells). Thus, the normalized flux of energy demand in ATP, denoted as JE(N) , representing the energy demand per unit of biomass, is proportional to the growth rate λ (see Appendix 2.1 for details): JE(N)=ηE⋅λ, (2)

where ηE is an energy coefficient (see Eqs. S25–S26 for details). By converting all energy currencies (such as NADH, FADH2, etc.) into ATP, the normalized energy fluxes for respiration and fermentation are given by Jr(E)=βr(A)⋅Jr(N)/2 and Jf(E)=βf(A)⋅Jf(N)/2 , where βr(A) and βf(A) are the stoichiometric coefficients of ATP production per glucose in each pathway (see Appendixfig. 1C-E and Appendix 2.1 for details). The denominator coefficient of “2” is derived from the stoichiometry of the coarse-grained reaction M1→2M2 (see Fig. 1A-B). Applying the criteria of flux balance (i.e., mass conservation; see Appendix 1.3) at each intermediate node (Mi, i = 1, …, 5) and precursor pool (Pool i, i = a1, a2, b, c, d), along with the constraints of proteome allocation (see Eq. 1) and energy demand (see Eq. 2), we obtain the relations between normalized energy fluxes and growth rate for a given nutrient condition with a fixed κA (see Appendix 2.1 for details): {Jr(E)+Jf(E)=φ⋅λ,Jr(E)εr+Jf(E)εf=ϕmax−ψ⋅λ, (3)

Earlier version: where φ is a constant coefficient primarily determined by ηE (see Eq. S33), and φ·λ represents the normalized flux of energy demand, excluding the energy biogenesis from the biomass synthesis pathway. The coefficients ψ, ψr and ψf are functions of κA · ψ-1 denotes the proteome efficiency for biomass generation in the biomass synthesis pathway (see Eq. S32), defined as ψ-1 = λ/ϕBM (see Appendix 2.1). ψr and ψf represent the proteome efficiencies for energy biogenesis in the respiration and fermentation pathways, respectively, defined as the normalized energy fluxes expressed in ATP generated per proteomic mass fraction, with εr=Jr(E)/ϕr and εf=Jf(E)/ϕf . Hence, {εr=βr(A)1/κA+1/κr(A),εr=βf(A)1/κA+1/κf(A), (4)

Revised version: where φ is a constant coefficient primarily determined by the coefficient ηE (see Eq. S33), and φ·λ represents the normalized flux of energy demand, excluding the energy biogenesis from the biomass synthesis pathway. The coefficients ψ, ψr and ψf are functions of κA, such that their values are highly dependent on nutrient conditions. ψ−1 denotes the proteome efficiency for biomass generation in the biomass synthesis pathway (see Eq. S32), defined as ψ-1 ≡ λ/ϕBM (see Appendix 2.1). ψr and ψf represent the proteome efficiencies for energy biogenesis in the respiration and fermentation pathways, respectively, defined as the normalized energy fluxes expressed in ATP generated per proteomic mass fraction, with εr≡Jr(E)/ϕr and εf≡Jf(E)/ϕf . Hence, {εr=βr(A)1/κA+1/κr(A),εf=βf(A)1/κA+1/κf(A), (4)

where both κr(A) and κf(A) are composite parameters that can be approximated as constants, with 1/κr(A) ≡ 1/κ1 + 2/κ2 + 2/κ3 + 2/κ4 and 1/κf(A) ≡ 1/κ1 + 2/κ2 + 2/κ6 (see Appendices 1.5 and 2.1 for details).

Origin of overflow metabolism

The standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) is exemplified by the experimental data (Basan et al., 2015) presented in Fig. 1C, where the fermentation flux exhibits a threshold-analog dependence on the growth rate λ. It is well established that respiration is significantly more efficient than fermentation in terms of energy biogenesis per unit of carbon (i.e., βr(A)>βf(A) ) (Nelson et al., 2008; Vander Heiden et al., 2009). Then, why do cells bother to use the seemingly wasteful fermentation pathway? We proceed to address this issue by applying optimal protein allocation (Scott et al., 2010; Wang et al., 2019) within the framework of optimal growth.

Earlier version: For cell proliferation in a given nutrient condition (i.e., with a fixed κA), the values of εr, εf and ψ are determined (Eqs. 4 and S32). However, the growth rate λ can be influenced by protein allocation between respiration and fermentation, specifically λr and λf, according to the governing equation (Eq. 3). If εr > εf, that is, if the proteome efficiency in respiration is higher than that in fermentation, then λ=ϕmax−Jr(E)(1/εf−1/εr)Ψ+φ/εr≤ϕmaxΨ+φ/εr . The optimal growth strategy is ϕf=Jf(E)=0 , meaning that the cell exclusively uses respiration. Conversely, if εf> εr, then ϕr=Jr(E)=0 is optimal, and the cell solely uses fermentation. In either case, the choice between respiration and fermentation for optimal growth is determined by weighing the proteome efficiencies.

Revised version: For cell proliferation in a given nutrient condition (i.e., with a fixed κA), the values of εr, εf and ψ are determined (Eqs. 4 and S32). However, the growth rate λ can be influenced by protein allocation between respiration and fermentation, specifically λr and λf, according to the governing equation (Eq. 3). If εr > εf, that is, if the proteome efficiency in respiration is higher than that in fermentation, then λ=ϕmax−Jr(E)(1/εf−1/εr)Ψ+φ/εr≤ϕmaxΨ+φ/εr . The optimal growth strategy is ϕf=Jf(E)=0 , meaning that the cell exclusively uses respiration. Conversely, if εf> εr, then ϕr=Jr(E)=0 is optimal, and the cell solely uses fermentation. In either case, the choice between respiration and fermentation for growth optimization is determined by comparing their proteome efficiencies.

Earlier version: In practice, both proteome efficiencies εr and εf are functions of nutrient quality κA, which can be influenced by the nutrient type and concentration of the carbon source (see Eqs. 4 and S27). Therefore, the choice for optimal growth may vary depending on the nutrient conditions. In nutrient-poor conditions where κA≪κr(A) and κA≪κf(A) , the proteome efficiencies can be approximated by εr≈βr(A)⋅κA and εf≈βf(A)⋅κA (see Eq. 4), and hence εr (κA)>εf(κA) (since βr(A)>βf(A) ), meaning that the proteome efficiency of respiration is higher than that of fermentation under nutrient-poor conditions. In rich media, however, using the parameters of κi derived from in vivo/in vitro experimental data of E. coli (see Appendix-tables 1–2 and Appendix 6.1–6.2), we obtain εr(κglucose(ST))<εf(κglucose(ST)) with Eq. 4 (see also Eqs. S39–S40), where (κglucose(ST)) represents the substrate quality of glucose at saturated concentration (abbreviated as “ST” in the superscript).

Revised version: In practice, both proteome efficiencies εr and εf are functions of nutrient quality κA, which can be significantly influenced by the nutrient type and concentration of the carbon source (see Eqs. 4 and S27). Therefore, the optimal growth strategy may vary depending on the nutrient conditions. In nutrient-poor conditions where κA≪κr(A) and κA≪κf(A) , the proteome efficiencies can be approximated by εr≈βr(A)⋅κA and εf≈βf(A)⋅κA (see Eq. 4), and hence εr (κA)>εf(κA) (since βr(A)>βf(A) ), meaning that the proteome efficiency of respiration is higher than that of fermentation under these conditions. In contrast, in rich media, using parameters for κi derived from in vivo/in vitro experimental data for E. coli (see Appendix-tables 1–2 and Appendix 6.1–6.2), we obtain εr(κglucose(ST))<εf(κglucose(ST)) with Eq. 4 (see also Eqs. S39–S40), where (κglucose(ST)) represents the substrate quality of glucose at saturated concentration (abbreviated as “ST” in the superscript).

Earlier version: This means that the proteome efficiency in fermentation is higher than that in respiration for bacteria in rich media. Indeed, recent studies have validated that the measured proteome efficiency in fermentation is higher than in respiration for E. coli in lactose at saturated concentration (Basan et al., 2015), i.e., εr(κlactose(ST))<εf(κlactose(ST)) . In Fig. 1E, we present the growth rate dependence of proteome efficiencies εr and εf in a three-dimensional (3D) form using the collected data shown in Appendix-table 1, where εr, εf and the growth rate λ all vary as functions of nutrient quality κA. Furthermore, the ratio Δ (defined as Δ(κA) ≡ εf (κA)/εr (κA)) is a monotonically increasing function of κA, and there exists a critical value of κA (denoted as κA(C) ; see Appendix 2.2 for details) satisfying Δ(κA(C))=1 . Below κA(C) , the cell grows slowly, and the choice for optimal growth is respiration (with εf < εr, λ = ϕmax · (ψ + φ/εr)-1), while above κA(C) , the cell grows faster, and the optimal growth strategy is fermentation (with εf < εr, λ = ϕmax · (ψ + φ/εf)-1)). The above analysis qualitatively explains the phenomenon of aerobic glycolysis.

Revised version: This indicates that the proteome efficiency in fermentation is higher than that in respiration for bacteria in rich media. Indeed, recent studies have validated that the measured proteome efficiency in fermentation is higher than in respiration for E. coli in lactose at saturated concentration (Basan et al., 2015), i.e., εr(κlactose(ST))<εf(κlactose(ST)) . In Fig. 1E, we present the growth rate dependence of proteome efficiencies εr and εf in a three-dimensional (3D) format using the collected data shown in Appendix-table 1, where εr, εf and the growth rate λ all vary as functions of nutrient quality κA. Furthermore, the ratio Δ (defined as Δ(κA) ≡ εf (κA)/εr (κA)) is a monotonically increasing function of κA, and there exists a critical value of κA (denoted as κA(C) ; see Appendix 2.2 for details) satisfying Δ(κA(C))=1 . Below κA(C) , where the nutrient is poorer and the cell grows slowly, the proteome efficiency of fermentation is lower than that of respiration (i.e., εf < εr), hence respiration is the optimal choice (with λ = ϕmax · (ψ + φ/εr)-1). Above κA(C) , the cell grows faster, and the optimal growth strategy is fermentation (with εf < εr, λ = ϕmax · (ψ + φ/εf)-1)). The above analysis qualitatively explains the phenomenon of aerobic glycolysis.

Earlier version: For a quantitative understanding of overflow metabolism, let us first consider the homogeneous case, where all cells share identical biochemical parameters. For optimal protein allocation, the relation between fermentation flux and growth rate is given by Jf(E)=φ⋅λ⋅θ(λ−λC) , where “Q” represents the Heaviside step function, and λC is the critical growth rate for a nutrient condition with κA(C) (i.e., λC ≡ λ( κA(C) )). Similarly, the growth rate dependence of respiration flux is Jf(E)=φ⋅λ⋅[ 1−θ(λ−λC) ] . These digital response outcomes are consistent with the numerical simulation findings of Molenaar et al. (Molenaar et al., 2009). However, they are clearly incompatible with the threshold-analog response observed in the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998).

Revised version: For a quantitative understanding of overflow metabolism, let us first consider the homogeneous case, where all cells share identical biochemical parameters. For optimal protein allocation, the relation between fermentation flux and growth rate under nutrient variation (with significantly varying κA) is given by Jf(E)=φ⋅λ⋅θ(λ−λC) , where “θ” represents the Heaviside step function, and λC denotes the critical growth rate corresponding to the nutrient condition with nutrient quality κA(C) (i.e., λC ≡ λ( κA(C) )). Similarly, the growth rate dependence of respiration flux is Jr(E)=φ⋅λ⋅[1−θ(λ−λC)] . These digital response outcomes are consistent with the numerical simulation findings of Molenaar et al. (Molenaar et al., 2009). However, they are clearly incompatible with the threshold-analog response observed in the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998).

Earlier version: To address this issue, we take into account cell heterogeneity, which is ubiquitous in both microbes (Ackermann, 2015; Bagamery et al., 2020; Balaban et al., 2004; Nikolic et al., 2013; Solopova et al., 2014; Wallden et al., 2016; Yaginuma et al., 2014; Zhang et al., 2018) and tumor cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016). In the context of the Warburg effect or overflow metabolism, experimental studies have reported significant metabolic heterogeneity in the choice between respiration and fermentation within a cell population (Bagamery et al., 2020; Duraj et al., 2021; Hensley et al., 2016; Nikolic et al., 2013). Motivated by the observation that the turnover number (kcat value) of a catalytic enzyme varies considerably between in vitro and in vivo measurements (Davidi et al., 2016; Garcia-Contreras et al., 2012), we note that the concentrations of potassium and phosphate, which vary from cell to cell, have a significant impact on the kcat values of metabolic enzymes (Garcia-Contreras et al., 2012). Therefore, within a cell population, there is a distribution of kcat values for a catalytic enzyme, commonly referred to as extrinsic noise (Elowitz et al., 2002). For simplicity, we assume that the kcat values for each enzyme follow a Gaussian distribution. This leads to varied distributions of single-cell growth rate in different carbon sources (see Eqs. S155–S157, S163–S165), which has been fully verified by recent experiments using isogenic E. coli at single-cell resolution (Wallden et al., 2016) (Appendix-fig. 2B). Consequently, the critical growth rate λC should follow a Gaussian distribution Ν(μλC,σλC2) within a cell population (see Appendix 7 for details), where μλC is approximated by the deterministic result of λC (Eq. S43). We assume that the coefficient of variation (CV) is σλC/μλC = 12%, or equivalently, the CV for the catalytic rate of each metabolic enzyme is 25%. Thus, we obtain the growth rate dependence of fermentation and respiration fluxes (see Appendix 2.3 for details): {Jf(N)(λ)=φ⋅λβf(A)⋅[erf(λ−μλC2σλC)+1],Jr(N)(λ)=φ⋅λβr(A)⋅[1−erf(λ−μλC2σλC)], (5)

Revised version: To address this issue, we take into account cell heterogeneity, which is ubiquitous in both microbes (Ackermann, 2015; Bagamery et al., 2020; Balaban et al., 2004; Nikolic et al., 2013; Solopova et al., 2014; Wallden et al., 2016; Yaginuma et al., 2014; Zhang et al., 2018) and tumor cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016; Shibao et al., 2018). In the context of the Warburg effect or overflow metabolism, experimental studies have reported significant metabolic heterogeneity in the choice between respiration and fermentation within a cell population (Bagamery et al., 2020; Duraj et al., 2021; Hensley et al., 2016; Nikolic et al., 2013). Motivated by the observation that the turnover number (kcat value) of a catalytic enzyme varies considerably between in vitro and in vivo measurements (Davidi et al., 2016; García-Contreras et al., 2012), we note that the concentrations of potassium and phosphate, which vary from cell to cell, have a significant impact on the kcat values of metabolic enzymes (García-Contreras et al., 2012). Therefore, within a cell population, there is a distribution of kcat values for a catalytic enzyme, commonly referred to as extrinsic noise (Elowitz et al., 2002). For simplicity, we assume that the kcat values for each enzyme follow a Gaussian distribution. Consequently, the proteome efficiencies εr and εf, which are crucial for determining the choice between respiration and fermentation, also follow Gaussian distributions (see Appendix 7 for details). This variability leads to diverse distributions of single-cell growth rates across different carbon sources (see Eqs. S155–S157, S163–S165), which has been fully verified by recent experiments using isogenic E. coli at single-cell resolution (Wallden et al., 2016) (Appendix-fig. 2B). Accordingly, the critical growth rate λC is expected to follow a Gaussian distribution Ν(μλC,σλC2) within a cell population (see Appendix 7 for details), where μλC is approximated by the deterministic result of λC (Eq. S43). Assuming the coefficient of variation (CV) of λC is σλC/μλC = 12%, or equivalently that the CV for the catalytic rate of each metabolic enzyme is 25%, we derive the growth rate dependence of fermentation and respiration fluxes (see Appendix 2.3 for details): {Jf(N)(λ)=φ⋅λβf(A)⋅[erf(λ−μλC2σλC)+1],Jr(N)(λ)=φ⋅λβr(A)⋅[1−erf(λ−μλC2σλC)], (5)

Earlier version: where “erf” represents the error function. The fermentation flux exhibits a threshold-analog relation with the growth rate (the red curves in Figs. 1C-D, 2B-C, and 3B, D, F), while the respiration flux (the blue curve in Fig. 1D) decreases with an increase in fermentation flux. In Fig. 1C-D, we observe that the model results (see Eq. 5 and Appendix 8; parameters are set based on the experimental data shown in Appendix-table S1) quantitatively agree with the experimental data from E. coli (Basan et al., 2015; Holms, 1996). The fermentation flux was determined by the acetate secretion rate Jactate(M)=2Jf(N) , and the respiration flux was deduced from the carbon dioxide flux JCO2,r(M)=6Jr(N) (the superscript “(M)” represents the measurable flux in the unit of mM/OD600/h; see Appendix 8.1 for details). Therefore, by incorporating cell heterogeneity, our model of optimal protein allocation quantitatively explains overflow metabolism.

Revised version: where “erf” represents the error function. The fermentation flux exhibits a threshold-analog relation with the growth rate (the red curves in Figs. 1C-D, 2B-C, and 3B, D, F), while the respiration flux (the blue curve in Fig. 1D) decreases as the fermentation flux increases. In Fig. 1C-D, we observe that the model results (see Eq. 5 and Appendix 8; parameters are set based on the experimental data shown in Appendix-table 1) quantitatively agree with the experimental data from E. coli (Basan et al., 2015; Holms, 1996). The fermentation flux is represented by the acetate secretion rate Jactate(M)=2Jf(N) , and the respiration flux is exemplified by the carbon dioxide flux JCO2,r(M)=6Jr(N) (the superscript “(M)” represents the measurable flux in the unit of mM/OD600/h; see Appendix 8.1 for details). By incorporating cell heterogeneity, our model of optimal protein allocation quantitatively explains overflow metabolism.

Influence of protein overexpression on overflow metabolism in E. coli.

(A) A 3D plot of the relations among fermentation flux, growth rate, and the expression level of useless proteins. In this plot, both the acetate excretion rate and growth rate vary as bivariate functions of the nutrient quality of a Group A carbon source (denoted as κA) and the useless protein expression encoded by lacZ gene (denoted as ϕZ perturbation; see Eqs. S57 and S160). (B) Growth rate dependence of the acetate excretion rate upon ϕZ perturbation for each fixed nutrient condition (see Eqs. S58 and S160). (C) Growth rate dependence of the acetate excretion rate as κA varies (see Eqs. S57 and S160), with each fixed expression level of LacZ.

Influence of energy dissipation, translation inhibition, and carbon source category alteration on overflow metabolism in E. coli.

(A) A 3D plot of the relations among fermentation flux, growth rate, and the energy dissipation coefficient (see Eqs. S70 and S160). (B) Growth rate dependence of the acetate excretion rate as the nutrient quality κA varies, with each fixed energy dissipation coefficient determined by or fitted from experimental data. (C) A 3D plot of the relations among fermentation flux, growth rate, and the translation efficiency (see Eqs. 85 and S160). Here, the translation efficiency is adjusted by the dose of chloramphenicol (Cm). (D) Growth rate dependence of the acetate excretion rate as κA varies, with each fixed dose of Cm. (E) Coarse-grained model for pyruvate utilization. (F) The growth rate dependence of fermentation flux in pyruvate (see Eqs. 105 and S160) significantly differs from that of the Group A carbon sources (see Eqs. 47 and S160).

Testing the model through perturbations

To further test our model, we systematically investigate its predictions under various types of perturbations and compare them with experimental data from existing studies (Basan et al., 2015; Holms, 1996) (see Appendices 3 and 4.1 for details).

First, we consider the proteomic perturbation caused by overexpression of useless proteins encoded by the lacZ gene (i.e., ϕZ perturbation) in E. coli. The net effect of the ϕZ perturbation is that the maximum fraction of the proteome available for resource allocation changes from ϕmax to ϕmax − ϕZ (Basan et al., 2015), where ϕZ is the proteomic mass fraction of useless proteins. In a cell population, the critical growth rate λC (ϕZ) still follows a Gaussian distribution N(μλC(ϕZ),σλC(ϕZ)2) , where the CV of λC (ϕZ) remains unchanged. Consequently, the growth rate dependence of fermentation flux changes to Jf(N)=φ⋅λβf(A)⋅[erf(λ−μλC(ϕZ)2σλC(ϕZ))+1] (see Appendix 3 for model perturbation results regarding respiration flux), where both the growth rate λ(κA, ϕZ) and the normalized fermentation flux Jf(N)(κA,ϕZ) are bivariate functions of κA and ϕZ (see Eqs. S49 and S56–S57). For each degree of LacZ expression (with fixed ϕZ), similar to wild-type strains, the fermentation flux exhibits a threshold-analog response to growth rate as κA varies (see Fig. 2C), which agrees quantitatively with experimental results (Basan et al., 2015). The shifts in the critical growth rate λC(ϕZ) are fully captured by μλC(ϕZ)=μλC(0)(1−ϕZ/ϕmax) . In contrast, for nutrient conditions with each fixed κA, since the growth rate changes with ϕZ just like λC (ϕZ): λ(κA, ϕZ) = λ (κA, 0)(1 − ϕZ/ϕmax), the fermentation flux is then proportional to the growth rate for the varying levels of LacZ expression: Jf(N)=φβf(A)⋅[erf(λ(κA)−μλC(0)2σλC(0))+1]⋅λ , where the slope is a monotonically increasing function of the substrate quality κA. These scaling relations are well validated by the experimental data (Basan et al., 2015) shown in Fig. 2B. Finally, in the case where both κA and ϕZ are free to vary, the growth rate dependence of fermentation flux presents a threshold-analog response surface in a 3D plot, where ϕZ appears explicitly as the y-axis (see Fig. 2A). Experimental data points (Basan et al., 2015) lie right on this surface, which is highly consistent with the model predictions.

Next, we study the influence of energy dissipation, which introduces an energy dissipation coefficient w to Eq. 2: JE(N)=ηE⋅λ+w . Similarly, the critical growth rate in this case, λC(w), follows a Gaussian distribution N(μλC(w),σλC(w)2) in a cell population. The relation between the growth rate and fermentation flux can be characterized by (see Appendix 3.2 for details): Jf(N)=φ⋅λ+wβf(A)⋅[erf(λ−μλC(w)2σλC(w)2)+1] . In Fig. 3A-B, we present a comparison between the model results and experimental data (Basan et al., 2015) in 3D and 2D plots, which demonstrate good agreement. A notable characteristic of energy dissipation, as distinguished from ϕZ perturbation, is that the fermentation flux increases despite a decrease in the growth rate when κA is fixed.

We proceed to analyze the impact of translation inhibition with different sub-lethal doses of chloramphenicol on E. coli. This type of perturbation introduces an inhibition coefficient ι to the translation rate, thus turning κt into κt/(ι + 1). Still, the critical growth rate λC(ι) follows a Gaussian distribution N(μλC(ι),σλC(ι)2) , and then, the growth rate dependence of fermentation flux is is given by: Jf(N)=φ⋅λβf(A)⋅[erf(λ−μλC(ι)2σλC(ι)2)+1] (see Appendix 3.3 for details). In Appendix-fig. 2D-E, we observe that the model predictions are generally consistent with the experimental data (Basan et al., 2015). However, a noticeable systematic discrepancy arises when the translation rate is low. Therefore, we consider maintenance energy, which is typically tiny and generally negligible for bacteria over the growth rate range of interest (Basan et al., 2015; Locasale and Cantley, 2010; Neidhardt, 1996). Encouragingly, by assigning a very small value to the maintenance energy coefficient w0 (w0 = 2.5 (h-1)), the model results for the growth rate-fermentation flux relation Jf(N)=φ⋅λ+w0βf(A)⋅[erf(λ−μλC(ι)2σλC(ι))+1] quantitatively agree with experiments (Basan et al., 2015) (see Fig. 3C-D and Appendix 3.3 for details).

Earlier version: Finally, we consider the alteration of nutrient categories by switching to a non-Group A carbon source: pyruvate, which enters the metabolic network from the endpoint of glycolysis (Neidhardt et al., 1990; Nelson et al., 2008). The coarse-grained model for pyruvate utilization is shown in Fig. 3E (see also Fig. 1A), which shares identical precursor pools with those for Group A carbon sources, yet has several differences in the coarse-grained reactions. The growth rate dependencies of both the proteome efficiencies (see Appendix-fig. 2H) and energy fluxes (see Fig. 3F) are qualitatively similar to those of Group A carbon source utilization, while there are quantitative differences in the coarse-grained parameters (see Appendices 4.1 and 8 for details). Most notably, the critical growth rate λC(py) and the ATP production per glucose in the fermentation pathway βf(py) for pyruvate utilization are noticeably smaller than those for Group A sources (i.e., λC and βf(A) , respectively). Consequently, the growth rate dependence of fermentation flux in pyruvate should present a distinctly different curve from that of Group A carbon sources (see Eqs. 5 and S105), which is fully validated by experimental results (Holms, 1996) (see Fig. 3F).

Revised version: Finally, we consider the alteration of nutrient categories by switching to a non-Group A carbon source: pyruvate, which enters the metabolic network from the endpoint of glycolysis (Neidhardt et al., 1990; Nelson et al., 2008). The coarse-grained model for pyruvate utilization is shown in Fig. 3E (see also Fig. 1A), which shares identical precursor pools with those for Group A carbon sources, yet has several differences in the coarse-grained reactions. The growth rate dependencies of both the proteome efficiencies (see Appendix-fig. 2H) and energy fluxes (see Fig. 3F) are qualitatively similar to those of Group A carbon source utilization, while there are quantitative differences in the coarse-grained parameters (see Appendices 4.1 and 8 for derivation details). Most notably, the critical growth rate λC(py) and the ATP production per glucose in the fermentation pathway βf(py) for pyruvate utilization are noticeably smaller than those for Group A sources (i.e., λC and βf(A) , respectively). Consequently, the growth rate dependence of fermentation flux in pyruvate should present a distinctly different curve from that of Group A carbon sources (see Eqs. 5 and S105), which is fully validated by experimental results (Holms, 1996) (see Fig. 3F).

Enzyme allocation under perturbations

As mentioned above, our coarse-grained model is topologically identical to the central metabolic network (Fig. 1A) and can thus predict enzyme allocation for each gene in glycolysis and the TCA cycle (see Appendix-fig. 1B and Appendix-table 1) under various types of perturbations. In Fig. 1B, the intermediate nodes M1, M2, M3, M4, and M5 represent G6P, PEP, acetyl-CoA, α-ketoglutarate, and oxaloacetate, respectively. Therefore, ϕ1 and ϕ2 correspond to enzymes involved in glycolysis (or at the junction of glycolysis and the TCA cycle), while ϕ3 and ϕ4 correspond to enzymes in the TCA cycle (see Fig. 1A-B and Appendix 2.1).

We first consider enzyme allocation under carbon limitation by varying the nutrient type and concentration of a Group A carbon source (i.e., κA perturbation). This has been extensively studied in more simplified models (Hui et al., 2015; You et al., 2013), where the growth rate dependence of enzyme allocation under κA perturbation is generally described by a C-line response (Hui et al., 2015; You et al., 2013). Specifically, the genes responsible for digesting carbon compounds exhibit a linear increase in gene expression as the growth rate decreases (Hui et al., 2015; You et al., 2013). However, when it comes to enzymes catalyzing reactions between intermediate nodes, we gathered experimental data from existing studies (Hui et al., 2015) and found that the enzymes in glycolysis exhibit a completely different response pattern compared to those in the TCA cycle (Appendix-fig. 3A-B). This discrepancy cannot be explained by the C-line response. To address this issue, we apply the coarse-grained model described above (see Fig. 1B) to calculate the growth rate dependence of enzyme allocation for each ϕi(i = 1, 2, 3, 4) using model settings for wild-type strains, with no fitting parameters influencing the shape (see Eqs. S118–S119 and Appendix 8). In Fig. 4A-B and Appendix-fig. 3C-D, we see that the model predictions overall match with the experimental data (Hui et al., 2015) for representative genes from either glycolysis or the TCA cycle, and maintenance energy (with w0 = 2.5 (h-1)) has a negligible effect on this process. Still, there are minor discrepancies that arise from the basal expression of metabolic genes, which may be attributed to the fact that our model deals with relatively stable growth conditions while microbes need to be prepared for fluctuating environments (Basan et al., 2020; Kussell and Leibler, 2005; Mori et al., 2017).

Relative protein expression of central metabolic enzymes in E. coli under nutrient limitation and proteomic perturbation.

(A, C) Relative protein expression of representative genes from glycolysis. (B, D) Relative protein expression of representative genes from the TCA cycle. (A, B) Results of the perturbation through changes in nutrient quality κA (see Eq. S119). (C, D) Results of proteomic perturbation via varied levels of expression of the useless protein LacZ (i.e., ϕZ perturbation; see Eq. S121).

We proceed to analyze the influence of ϕZ perturbation and energy dissipation. In both cases, our model predicts a linear response to growth rate reduction for all genes in either glycolysis or the TCA cycle (see Appendix 5.2–5.3 for details). For ϕZ perturbation, all predicted slopes are positive, and there are no fitting parameters involved (Eqs. S120–S121). In Fig. 4C-D and Appendix-fig. 3E-J, we show that our model quantitatively illustrates the experimental data (Basan et al., 2015) for representative genes in the central metabolic network, and there is better agreement with experiments (Basan et al., 2015) by incorporating the maintenance energy (with w0 = 2.5 (h-1) as aforementioned). For energy dissipation, however, the predicted slopes of the enzymes corresponding to ϕ4 are negative, and there is a constraint that the slope signs of the enzymes corresponding to the same ϕi(i = 1, 2, 3) should be the same. In Appendix-fig. 3K-N, we see that the model results (Eqs. S127 and S123) are consistent with experiments (Basan et al., 2015).

Explanation of the Crabtree effect in yeast and the Warburg effect in cancer cells

Earlier version: We proceed to apply our model to explain the Crabtree effect in yeast (Bagamery et al., 2020; De Deken, 1966; Shen et al., 2024) and the Warburg effect in tumors (Bartman et al., 2023; Duraj et al., 2021; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Shen et al., 2024; Vander Heiden et al., 2009) with slight modifications using the optimal growth principle combined with cell heterogeneity (see Appendix 9 and Appendix-fig. 5). For yeast and tumors, similar to the case of E. coli, the proteome efficiencies εr and εf are both increasing functions of the nutrient quality κA (see Eq. S170). Specifically, under poor nutrient conditions (i.e., εA is small), the proteome efficiency in respiration is higher than that in fermentation: εr > εf (see Eqs. S174–S175), making respiration the best choice for optimal growth (see Eq. S171). Conversely, when the nutrient is abundant and εf > εr, aerobic glycolysis (i.e., fermentation) becomes the optimal growth strategy (see Eq. S172). Further combination with cell heterogeneity results in the standard picture of overflow metabolism, which has indeed been observed in yeast (van Hoek et al., 1998). However, it remains challenging to tune the growth rate of cancer cells in vivo.

Revised version: We proceed to apply our model to explain the Crabtree effect in yeast (Bagamery et al., 2020; De Deken, 1966; Shen et al., 2024) and the Warburg effect in tumors (Bartman et al., 2023; Duraj et al., 2021; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Shen et al., 2024; Vander Heiden et al., 2009) with slight modifications using the optimal growth principle combined with cell heterogeneity (see Appendix 9 and Appendix-fig. 5). For yeast and tumors, similar to the case of E. coli, the proteome efficiencies εr and εf are both increasing functions of nutrient quality κA (see Eq. S170). Under poor nutrient conditions (i.e., εA is small), the proteome efficiency in respiration is higher than that in fermentation: εr > εf (see Eqs. S174–S175), making respiration the optimal choice for growth optimization (see Eq. S171). Conversely, when the nutrients are abundant and εf > εr, aerobic glycolysis (i.e., fermentation) becomes the optimal growth strategy (see Eq. S172). Further combination with cell heterogeneity results in the standard picture of overflow metabolism, which has indeed been observed in yeast (van Hoek et al., 1998). However, it remains challenging to tune the growth rate of cancer cells in vivo.

Recently, Shen et al. (Shen et al., 2024) discovered that the proteome efficiency measured at the cell population level in respiration (i.e., 〈εr〉; where “〈 〉” denotes the population average) is higher than that in fermentation (i.e., 〈εf〉) for many yeast and cancer cells, despite the presence of fermentation fluxes through aerobic glycolysis. Evidently, this finding (Shen et al., 2024) contradicts prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019), which hold that overflow metabolism arises because the proteome efficiency in fermentation is consistently higher than in respiration. Nevertheless, our model may resolve this puzzle due to the incorporation of two important features. First, our model predicts that the proteome efficiency in respiration is larger than that in fermentation when nutrient quality is low (see Eqs. S174–S175). Second, and crucially, by accounting for cell heterogeneity, our model allows a proportion of cells to have a higher proteome efficiency in fermentation than in respiration, even when the overall proteome efficiency in respiration at the cell population level is greater than that in fermentation (i.e., 〈εr〉 > 〈εf〉).

Earlier version: To compare our model results quantitatively with experimental data on yeast and tumors (Shen et al., 2024), we define Prf≡Jf(E)Jf(E)+Jr(E) as the fraction of ATP produced through fermentation. To account for cell heterogeneity, we apply Gaussian distributions to enzyme turnover numbers, as mentioned above. This yields the relationship between Prf(i.e., Jf(E)Jf(E)+Jr(E)) and 〈εr〉 and 〈εf〉 through derivations (see Eqs. S180–S190 and Appendix 9 for details): Jf(E)Jf(E)+Jr(E)=12[erf(1−〈εr〉/〈εf〉2⋅χεr2+χεf2⋅(〈εr〉/〈εf〉)2)+1], (6)

Revised version: To compare our model results quantitatively with experimental data on yeast and tumors (Shen et al., 2024), we define Prf≡Jf(E)Jf(E)+Jr(E) as the fraction of ATP produced through fermentation. To account for cell heterogeneity, we apply Gaussian distributions to enzyme turnover numbers, as described above. This yields the relationship between Prf(i.e., Jf(E)Jf(E)+Jr(E)) and 〈εr〉 and 〈εf〉 through derivations (see Eqs. S180–S190 and Appendix 9 for details): Jf(E)Jf(E)+Jr(E)=12[erf(1−〈εr〉/〈εf〉2⋅χεr2+χεf2⋅(〈εr〉/〈εf〉)2)+1], (6)

where χεr and χεf represent the CVs of proteome efficiencies εr and εf, respectively. Due to the higher levels of cell heterogeneity in yeast (Bagamery et al., 2020) and cancer cells (Duraj et al., 2021; Hanahan and Weinberg, 2011; Hensley et al., 2016), the CVs of εr and εf (i.e., χεr and χεf) in these cells are expected to be significantly higher than those in E. coli, although their precise values are unknown. The values for the variables shown in Eq. 6 can be obtained from experiments. Therefore, we plot the theoretical results from Eq. 6 using χεr and χεf values of 0.25, 0.40, and 0.58 to compare with experimental data from yeast and in vivo mouse tumors (Bartman et al., 2023; Shen et al., 2024). As shown in Fig. 5A-B, the theoretical results with χεr = χεf = 0.58 align quantitatively with the experimental data (Bartman et al., 2023; Shen et al., 2024) on both logarithmic and linear scales, demonstrating that our model has the potential to quantitatively explain the Crabtree effect in yeast and the Warburg effect in cancer cells.

Model comparison with data on the Crabtree effect in yeast and the Warburg effect in tumors.

(A) A linear scale representation on the y-axis. (B) A log scale representation on the y-axis. In (A-B), 〈εr〉 and 〈εf〉 represent the population averages of εr and εf, while χεr and χεf are the coefficients of variation (CVs) of εr and εf · 〈εr〉/〈εf〉 represents the ratio of proteome efficiency between respiration and fermentation at the population-averaged level, while Jf(E)/(Jf(E)+Jr(E)) stands for the fraction of energy flux generated by the fermentation pathway (see Eq. 6). The data for yeast in batch culture and chemostat were calculated from experimental data of S. cerevisiae and I. orientalis (Shen et al., 2024). The data for mouse tumors were calculated from in vivo experimental data of pancreatic ductal adenocarcinoma (PDAC) and leukemic spleen of mice (Bartman et al., 2023; Shen et al., 2024). See Appendix 10 for detailed information on the experimental data sources (Bartman et al., 2023; Shen et al., 2024).

Discussion

Earlier version: The phenomenon of overflow metabolism, or the Warburg effect, has been a long-standing puzzle in cell metabolism. Although many rationales have been proposed over the past century (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011), contradictions persist (Shen et al., 2024), leaving the origin and function of this phenomenon unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). In this study, we use E. coli as a typical example and demonstrate that overflow metabolism can be understood through optimal protein allocation combined with cell heterogeneity. Under nutrient-poor conditions, the proteome efficiency of respiration is higher than that of fermentation (see Fig. 1E), and thus the cell uses respiration to optimize growth. In rich media, however, the proteome efficiency of fermentation increases more rapidly and surpasses that of respiration (see Fig. 1E), leading the cell to adopt fermentation as the optimal growth strategy. In further combination with cell heterogeneity in enzyme catalytic rates (Davidi et al., 2016; Garcia-Contreras et al., 2012), our model quantitatively illustrates the threshold-analog response (Basan et al., 2015; Holms, 1996) in overflow metabolism (see Fig. 1C). Furthermore, it quantitatively explains the data on the Crabtree effect in yeast and the Warburg effect in cancer cells (Bartman et al., 2023; Shen et al., 2024).

Revised version: The phenomenon of overflow metabolism, or the Warburg effect, has been a long-standing puzzle in cell metabolism. Although many rationales have been proposed over the past century (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Molenaar et al., 2009; Niebel et al., 2019; Peebo et al., 2015; Pfeiffer et al., 2001; Shlomi et al., 2011; Vander Heiden et al., 2009; Varma and Palsson, 1994; Vazquez et al., 2010; Vazquez and Oltvai, 2016; Zhuang et al., 2011), contradictions persist (Shen et al., 2024), leaving the origin and function of this phenomenon unclear (DeBerardinis and Chandel, 2020; Hanahan and Weinberg, 2011; Liberti and Locasale, 2016; Vander Heiden et al., 2009). In this study, we use E. coli as a typical example and demonstrate that overflow metabolism can be understood through optimal protein allocation combined with cell heterogeneity. Under nutrient-poor conditions, the proteome efficiency of respiration is higher than that of fermentation (see Fig. 1E), and thus the cell uses respiration to optimize growth. In rich media, however, the proteome efficiency of fermentation increases more rapidly and surpasses that of respiration (see Fig. 1E), leading the cell to adopt fermentation as the optimal growth strategy. In further combination with cell heterogeneity in enzyme catalytic rates (Davidi et al., 2016; García-Contreras et al., 2012), our model quantitatively illustrates the threshold-analog response (Basan et al., 2015; Holms, 1996) in overflow metabolism (see Fig. 1C). Furthermore, it quantitatively explains the data on the Crabtree effect in yeast and the Warburg effect in cancer cells (Bartman et al., 2023; Shen et al., 2024).

Added in the revised version: Mechanistically, the optimal growth strategy for the binary choice between respiration and fermentation can be facilitated by the direct sensing and comparison of proteome efficiencies between the two pathways (see Appendix 2.4). A growing body of evidence suggests that the cyclic AMP (cAMP)-cAMP receptor protein (CRP) system plays a crucial role in sensing proteome efficiency and executing the optimal strategy (Basan et al., 2015; Towbin et al., 2017; Valgepea et al., 2010; Wehrens et al., 2023). However, it has also been suggested that the cAMP-CRP system alone is insufficient, and that additional regulators remain to be identified to fully elucidate this mechanism (Basan et al., 2015; Valgepea et al., 2010). Furthermore, since the binary choice between respiration and fermentation is driven by the comparison of proteome efficiencies, the optimal growth principle in our model can be relaxed to the case where efficient protein allocation is required only for enzymes, rather than ribosomes. This allows our model to remain applicable under suboptimal growth conditions (see Appendix 2.4 for details), where recent experimental studies have shown that the inactive portion of ribosomes (i.e., ribosomes not bound to mRNAs) may vary with culturing conditions (Dai et al., 2017; Li et al., 2018) and between individual cells within the same culture (Pavlou et al., 2025), despite an overall trend toward growth optimization.

Earlier version: In existing rationales (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez and Oltvai, 2016), the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) has primarily been illustrated by a threshold-linear response, which heavily depends on the assumption that cells optimize their growth rate based on a given rate of carbon influx under each nutrient condition (or similar equivalents, see Appendix 6.3). However, in practice, for microbes or tumor cells grown in vitro or in vivo, the given factors are the identity and concentration of the nutrient (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the rate of carbon influx. Additionally, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) suggest that overflow metabolism originates from the proteome efficiency in fermentation always being higher than that in respiration (see Appendix 6.3 for details). While it has been observed in E. coli that proteome efficiency in fermentation is higher than that in respiration for cells cultured in lactose at saturated concentration (Basan et al., 2015), Shen et al. (Shen et al., 2024) reported that for many yeast and cancer cells, the proteome efficiency in fermentation is noticeably lower than that in respiration, despite the presence of aerobic glycolytic fermentation flux. This observation (Shen et al., 2024) evidently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Our model resolves this puzzle by significantly differing from existing rationales in its optimization principle, where we purely optimize cell growth rate without imposing a special constraint on carbon influx (see Appendix 6.3 for details). More importantly, our model incorporates cell heterogeneity, which is crucial for both explaining the threshold-analog response in overflow metabolism and for resolving this puzzle raised by Shen et al. (Shen et al., 2024).

Revised version: In existing rationales (Basan et al., 2015; Chen and Nielsen, 2019; Majewski and Domach, 1990; Shlomi et al., 2011; Varma and Palsson, 1994; Vazquez and Oltvai, 2016), the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998) has primarily been illustrated by a threshold-linear response, which largely relies on the assumption that cells optimize their growth rate for a given rate of carbon influx under each nutrient condition (or similar equivalents, see Appendix 6.3). However, in practice, for microbes or tumor cells grown in vitro or in vivo, the given factors are the identity and concentration of the nutrient (Molenaar et al., 2009; Scott et al., 2010; Wang et al., 2019), rather than the rate of carbon influx. Additionally, prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019) suggest that overflow metabolism originates from the proteome efficiency in fermentation always being higher than that in respiration (see Appendix 6.3 for details). While it has been observed in E. coli that proteome efficiency in fermentation is higher than that in respiration for cells cultured in lactose at saturated concentration (Basan et al., 2015), Shen et al. (Shen et al., 2024) reported that for many yeast and cancer cells, the proteome efficiency in fermentation is noticeably lower than that in respiration, despite the presence of aerobic glycolytic fermentation flux. This observation (Shen et al., 2024) evidently contradicts the prevalent explanations (Basan et al., 2015; Chen and Nielsen, 2019). Our model resolves this puzzle by significantly differing from existing rationales in its optimization principle, where we purely optimize cell growth rate purely through protein allocation without imposing a special constraint on carbon influx (see Appendix 6.3 for details). More importantly, our model incorporates cell heterogeneity, which is crucial for both explaining the threshold-analog response in overflow metabolism and for resolving this puzzle raised by Shen et al. (Shen et al., 2024).

Earlier version: In the homogeneous case, the optimal growth strategy for growth rate dependent fermentation flux results in a digital response (see Eq. S44), corresponding to an elementary flux mode (Müller et al., 2014; Wortel et al., 2014), which aligns with the numerical study by Molenaar et al. (Molenaar et al., 2009), yet is incompatible with the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998). Furthermore, cells would not generate fermentation flux if the proteome efficiency in fermentation were higher than that in respiration under the optimal growth framework. By incorporating heterogeneity in enzyme catalytic rates (Davidi et al., 2016; Garcia-Contreras et al., 2012), the critical growth rate (i.e., threshold) shifts from a single value to a Gaussian distribution (see Eq. S45 and Appendix 7 for details; see also Appendix-fig. 4) across a cell population, thereby turning a digital response into the threshold-analog response observed in overflow metabolism (see Fig. 1C). Moreover, cell heterogeneity allows a fraction of cells to possess a larger proteome efficiency in fermentation than in respiration despite the overall proteome efficiency in respiration at the cell population level is higher than in fermentation. This mechanism facilitates the fermentation flux in yeast and cancer cells observed by Shen et al. (Shen et al., 2024) (see Fig. 5A-B).

Revised version: In the homogeneous case, the optimal growth strategy for growth rate dependent fermentation flux results in a digital response (see Eq. S44), corresponding to an elementary flux mode (Müller et al., 2014; Wortel et al., 2014), which aligns with the numerical study by Molenaar et al. (Molenaar et al., 2009), but is incompatible with the standard picture of overflow metabolism (Basan et al., 2015; Holms, 1996; Meyer et al., 1984; Nanchen et al., 2006; van Hoek et al., 1998). Furthermore, in this case, cells would not generate fermentation flux if the proteome efficiency in fermentation were lower than that in respiration, under the optimal growth framework. By incorporating heterogeneity in enzyme catalytic rates (Davidi et al., 2016; García-Contreras et al., 2012), the critical growth rate (i.e., threshold) shifts from a single value to a Gaussian distribution (see Eq. S45 and Appendix 7 for details; see also Appendix-fig. 4) across a cell population, thereby turning a digital response into the threshold-analog response observed in overflow metabolism (see Fig. 1C). Moreover, cell heterogeneity allows a fraction of cells to possess a larger proteome efficiency in fermentation than in respiration despite the overall proteome efficiency in respiration at the cell population level is higher than in fermentation. This mechanism facilitates the fermentation flux in yeast and cancer cells observed by Shen et al. (Shen et al., 2024) (see Fig. 5A-B).

Earlier version: Our model results, based on cell heterogeneity, are further supported by observed distributions of single-cell growth rates in E. coli (Wallden et al., 2016) (see Appendix-fig. 2B), and by experiments involving various types of perturbations (Basan et al., 2015; Holms, 1996; Hui et al., 2015), both in terms of acetate secretion patterns and gene expression in the central metabolic network (see Fig. 2–4, Appendix-figs. 2D-E and 3). Furthermore, the heterogeneity patterns predicted by our model for fermentation and respiration modes within a cell population are validated by experiments on intra-tumor heterogeneity in glioblastoma using optical microscopy (Duraj et al., 2021) and align with the non-genetic heterogeneity observed in experiments with S. cerevisiae (Bagamery et al., 2020) and E. coli (Nikolic et al., 2013). Finally, our model can be broadly applied to address heterogeneity-related challenges in metabolism on a quantitative basis, including diverse metabolic strategies of cells in various environments (Bagamery et al., 2020; Duraj et al., 2021; Escalante-Chong et al., 2015; Hensley et al., 2016; Liu et al., 2015; Solopova et al., 2014; Wang et al., 2019).

Revised version: Our model results, based on cell heterogeneity, are further supported by observed distributions of single-cell growth rates in E. coli (Wallden et al., 2016) (see Appendix-fig. 2B), as well as by experiments involving various types of perturbations (Basan et al., 2015; Holms, 1996; Hui et al., 2015), both in terms of acetate secretion patterns and gene expression in the central metabolic network (see Fig. 2–4, Appendix-figs. 2D-E and 3). Furthermore, the heterogeneity patterns predicted by our model for fermentation and respiration modes in an isogenic cell population under the same culturing conditions are highly consistent with the non-genetic heterogeneity observed in single-cell experiments with E. coli (Nikolic et al., 2013) and S. cerevisiae (Bagamery et al., 2020, and align with experiments on intra-tumor heterogeneity in glioblastoma (Duraj et al., 2021; Shibao et al., 2018). Finally, our model can be broadly applied to address heterogeneity-related challenges in metabolism on a quantitative basis, including diverse metabolic strategies of cells in various environments (Bagamery et al., 2020; Duraj et al., 2021; Escalante-Chong et al., 2015; Hensley et al., 2016; Liu et al., 2015; Solopova et al., 2014; Wang et al., 2019).